/*
 * This exercise will focus on relational operators.
 * Fix the if statements so that the code correctly compares the values of a and b.
 *
 *
 */

// ❌ I AM NOT DONE

#include <assert.h>
#include <stdbool.h>
#include <stdio.h>


int main(void) {
	int a = 3;
	int b = 5;

	if (a b) {
		printf("equal\n");
	} else if (a  b) {
		printf("not equal\n");
	} else if (a  b) {
		printf("less than\n");
	} else if (a  b) {
		printf("greater than\n");
	} else if (a  b) {
		printf("less than or equal\n");
	} else if (a  b) {
		printf("greater than or equal\n");
	}

	return 0;
}
