/* 
 * Variables in C must be declared before they are used. Typically you will see them declared at the top of a function.
 *
 * Take a look at the program below and figure out what primitive datatype each variable should be and assign them accordingly.
*/

// ❌ I AM NOT DONE

int main() {
	x = 10;
	y = 1.0;
	z = 'A';

	return 0;
}
