/*
 * Variable naming seems pretty simple however there are some gotchas. 
 *
 * The primary gotchas are: "reserved words" and "illegal characters".
 *
 * Take a look at the below exercise and identify the variables that are 
 * causing the compilation error.
*/

// ❌ I AM NOT DONE

int main() {
  // [TODO] Identify the variable(s) that is causing the compilation error. 
  char foo1 = 'a';
  float float = 1.0;
  int bar = 1;
  char 1baz = 'a';
  
	return 0;
}

